```python
# Initialize Otter
import otter
grader = otter.Notebook("p6.ipynb")
```

```python
import p6_test
```

# Project 6: Airbnb

## Learning Objectives:

In this project, you will demonstrate how to:

* access and utilize data in CSV files,
* process real world datasets,
* use string methods and sorting function / method to order data.

Please go through [lab-p6](https://git.doit.wisc.edu/cdis/cs/courses/cs220/cs220-f22-projects/-/tree/main/lab-p6) before working on this project. The lab introduces some useful techniques related to this project.

## Testing your code:

Along with this notebook, you must have downloaded the file `p6_test.py`. If you are curious about how we test your code, you can explore this file, and specifically the value of the variable `expected_json`, to understand the expected answers to the questions.

## Project Description:

Data Science can help us understand user behavior on online platform services. This project is about the rooms listed on Airbnb. Since 2008, guests and hosts have used Airbnb to expand on traveling possibilities and present a more unique, personalized way of experiencing the world. `airbnb.csv` has data of nearly 50,000 listings on Airbnb from New York City, NY from the year 2019. This file includes a lot of information about the hosts, geographical availability of the listings, and other necessary metrics to make predictions and draw conclusions. You will be using various string manipulation methods that come with Python as well as creating some of your own functions to solve the problems posed. Happy coding!

## Dataset:

A small portion of the dataset `airbnb.csv` you will be working with for this project is reproduced here:

room_id|name|host_id|host_name|neighborhood_group|neighborhood|latitude|longitude|room_type|price|minimum_nights|number_of_reviews|last_review|reviews_per_month|calculated_host_listings_count|availability_365
------|------|------|------|------|------|------|------|------|------|------|------|------|------|------|------|
2539|Clean & quiet apt home by the park|2787|John|Brooklyn|Kensington|40.64749000000001|-73.97237|Private room|149|1|9|2018-10-19|0.21|6|365
2595|Skylit Midtown Castle|2845|Jennifer|Manhattan|Midtown|40.75362|-73.98376999999998|Entire home/apt|225|1|45|2019-05-21|0.38|2|355
3647|THE VILLAGE OF HARLEM....NEW YORK !|4632|Elisabeth|Manhattan|Harlem|40.80902|-73.9419|Private room|150|3|0|||1|365
3831|Cozy Entire Floor of Brownstone|4869|LisaRoxanne|Brooklyn|Clinton Hill|40.68514|-73.95976|Entire home/apt|89|1|270|2019-07-05|4.64|1|194
5022|Entire Apt: Spacious Studio/Loft by central park|7192|Laura|Manhattan|East Harlem|40.79851|-73.94399|Entire home/apt|80|10|9|2018-11-19|0.1|1|0
5099|Large Cozy 1 BR Apartment In Midtown East|7322|Chris|Manhattan|Murray Hill|40.74767|-73.975|Entire home/apt|200|3|74|2019-06-22|0.59|1|129

You can find more details on the dataset in [lab-p6](https://git.doit.wisc.edu/cdis/cs/courses/cs220/cs220-f22-projects/-/tree/main/lab-p6).

## Questions and Functions:

Let us start by importing all the modules we will need for this project.

```python
# it is considered a good coding practice to place all import statements at the top of the notebook
# please place all your import statements in this cell if you need to import any more modules for this project
import csv
```

#### Now, copy and paste the `process_csv` and `cell` functions from your [lab-p6](https://git.doit.wisc.edu/cdis/cs/courses/cs220/cs220-f22-projects/-/tree/main/lab-p6) notebook to the cell below.

You are expected to call the `process_csv` function correctly, and read the data on `airbnb.csv`. After reading the file, define the `csv_header`, and `csv_rows` variables as in [lab-p6](https://git.doit.wisc.edu/cdis/cs/courses/cs220/cs220-f22-projects/-/tree/main/lab-p6), and define the `cell` function.

**Important:** You **must** only use the `cell` function to extract data from the dataset. If you extract any data without explicitly using this function, you will **lose points** during manual review. It is recommended but **optional** that you use the `cell_v2` function defined in lab-p6. However, you **must** rename the function to `cell` in your p6 notebook. 

```python
def process_csv(filename):
    example_file = open(filename, encoding="utf-8")
    example_reader = csv.reader(example_file)
    example_data = list(example_reader)
    example_file.close()
    return example_data

csv_data = process_csv("airbnb.csv")

def cell(row_idx, col_name):
    col_idx = csv_header.index(col_name)
    val = csv_rows[row_idx][col_idx]
    if val == "":
        return None
    elif col_name == 'price' or col_name == 'number_of_reviews' or col_name == 'calculated_host_listings_count' or col_name == 'availability_365':
        val = int(val)
    elif col_name == 'reviews_per_month' or col_name == 'latitude' or col_name == 'longitude':
        val = float(val)
    elif col_name == 'minimum_nights':
        val = int(val)
    return val

csv_header = csv_data[0]
csv_header

csv_rows = csv_data[1:]
csv_rows

num_rows = len(csv_rows)
```

**Question 1:** What **unique** `neighborhood_groups` are included in the dataset?

Your output **must** be a *list* which stores all the **unique** neighborhood groups (i.e., without any duplicates). The order **does not** matter.

```python
neighborhood_groups = []
for idx in range(len(csv_rows)):
    neighborhood_groups.append(cell(idx, "neighborhood_group"))
set_betw = set(neighborhood_groups)
neighborhood_groups = list(set_betw)
neighborhood_groups
```

Output:
```
['Manhattan', 'Bronx', 'Staten Island', 'Queens', 'Brooklyn']
```

```python
grader.check("q1")
```

Output:
```
q1 results: All test cases passed!
```

**Question 2:** What is the **average** `price` of rooms in the dataset?

```python
# compute and store the answer in the variable 'avg_price', then display it
total_p = 0
count = 0
for idx in range(len(csv_rows)):
    total_p += cell(idx, "price")
    count += 1
avg_price = total_p/count
avg_price
```

Output:
```
152.7206871868289
```

```python
grader.check("q2")
```

Output:
```
q2 results: All test cases passed!
```

**Question 3:** How many rooms are in the `neighborhood` of *SoHo*?

```python
# compute and store the answer in the variable 'count_soho', then display it
count_soho = 0
for idx in range(len(csv_rows)):
    if cell(idx, "neighborhood") == "SoHo":
        count_soho += 1
count_soho
```

Output:
```
358
```

```python
grader.check("q3")
```

Output:
```
q3 results: All test cases passed!
```

### Function 1: `find_room_names(phrase)`

We require you to complete the below function to answer the next several questions (this is a **requirement**, and you will **lose points** if you do not implement this function).

```python
def find_room_names(phrase):
    
    list_new = []
    phrase = phrase.lower()
    for idx in range(num_rows):
        if cell(idx, "name") == None:
            continue
        if phrase in (cell(idx, "name")).lower():
            list_new.append(cell(idx, "name"))
            
    return list_new
```

**Question 4:** Find all room `names` that contain the string `"CBG"`.
    
Your output **must** be a *list*. The order **does not** matter. You **must** call the `find_room_names` function to answer this question.

```python
# compute and store the answer in the variable 'rooms_contain_cbg', then display it
rooms_contain_cbg = find_room_names("CBG")
rooms_contain_cbg
```

Output:
```
['CBG CtyBGd HelpsHaiti rm#1:1-4',
 'CBG Helps Haiti Room#2.5',
 'CBG Helps Haiti Rm #2',
 'CBG# 4Tiny room w/ huge window/AC',
 'CBG Helps Haiti Rm #3',
 'CBG HelpsHaiti #5 Suite']
```

```python
grader.check("q4")
```

Output:
```
q4 results: All test cases passed!
```

**Question 5:** Find all room `names` that contain **either** `"cinema"` **or** `"film"`.

Your output **must** be a *list*. The order **does not** matter, but if a room's `name` contains **both** `"cinema"` and `"film"`, then the room must be included **only once** in your list. You **must** call the `find_room_names` function to answer this question.

```python
# compute and store the answer in the variable 'rooms_contain_cinema_film', then display it
list_new = find_room_names("cinema") + find_room_names("film")
rooms_contain_cinema_film = list(set(list_new))
rooms_contain_cinema_film
```

Output:
```
['TV-PHOTO-FILM-CINEMA-ART GALLERY-MUSIC STUDIO-LOFT',
 'Brooklyn townhouse for filming',
 'Music Recording Mixing Film Photography Art',
 'Shoot. Film. Sleep. Unique Loft Space in Brooklyn.',
 'Clean music/film themed bedroom',
 "Downtown Filmmaker's Loft by WTC",
 'Film / photography location in unique apartment',
 'The Otheroom Bar/Event/Filming Space -read details',
 'Film Location',
 'Modern Townhouse for Photo, Film &  Daytime Events',
 'Cinema + gym included with room',
 'Cool apartment in Brooklyn with free cinema & gym',
 'Victorian Film location',
 'WoodyAllen FilmSet-Like Digs (Room)',
 'WoodyAllen FilmSet-Like Digs (Apt)',
 'Sunny private room featured in film',
 'HUGE LUX 2FLOOR 2 BDRMSOHO LOFTw/HOME CINEMA',
 'Stunning Chelsea 1BR w/ Gym, W/D, Doorman, Sundeck, Cinema, by Blueground',
 'Premium Chelsea 1BR w/ Gym, W/D, Doorman, Sundeck, Cinema, by Blueground',
 'Cinema Studio on Duplex Apt.']
```

```python
grader.check("q5")
```

Output:
```
q5 results: All test cases passed!
```

**Question 6:** Find the **longest** room `name` that contains the word `"fun"`.

There is a **unique** such room with the longest `name`, so you **do not** have to worry about breaking ties. You **must** call the `find_room_names` function to answer this question.

```python
# compute and store the answer in the variable 'funnest_room', then display it
funnest_room = ""
fun_names = find_room_names("fun")
for idx in range(len(fun_names)):
    if len(fun_names[idx]) > len(funnest_room):
        funnest_room = fun_names[idx]
funnest_room
```

Output:
```
'Homey 1BR in Fun, Central West Village w/ Doorman by Blueground'
```

```python
grader.check("q6")
```

Output:
```
q6 results: All test cases passed!
```

**Question 7:** Find all the `names` of the rooms which have `price` *0* and have **more than** *30* `reviews`.

Your output **must** be a *list*. The names **must** be sorted in **ascending (alphabetical) order**.

```python
# compute and store the answer in the variable 'no_cost_rooms', then display it
no_cost_rooms = []
for idx in range(num_rows):
    if cell(idx, "price") == 0 and cell(idx, "number_of_reviews") > 30:
        no_cost_rooms.append(cell(idx, "name"))
no_cost_rooms = sorted(no_cost_rooms)
no_cost_rooms
```

Output:
```
['Contemporary bedroom in brownstone with nice view',
 'Cozy yet spacious private brownstone bedroom',
 'Spacious comfortable master bedroom with nice view',
 '★Hostel Style Room | Ideal Traveling Buddies★']
```

```python
grader.check("q7")
```

Output:
```
q7 results: All test cases passed!
```

**Question 8:**  What `neighborhoods` are the rooms that have `prices` greater than *9999* located in?

Your output **must** be a *list* of **unique** neighborhoods (i.e., without any duplicates). The names **must** be sorted in **descending (reverse-alphabetical) order**.

```python
# compute and store the answer in the variable 'pricey_neighborhoods', then display it
pricey_neighborhoods = []
for idx in range(num_rows):
    if cell(idx, "price") > 9999:
        pricey_neighborhoods.append(cell(idx, "neighborhood"))
pricey_neighborhoods = list(set(pricey_neighborhoods))
pricey_neighborhoods.sort(reverse = True)
pricey_neighborhoods
```

Output:
```
['Upper West Side', 'Greenpoint', 'Astoria']
```

```python
grader.check("q8")
```

Output:
```
q8 results: All test cases passed!
```

**Question 9:** How many rooms received their `last_review` in or before *2015*?

You should **ignore** rooms for which the `last_review` data is missing.

**Hint:** You can find the date of the last review in the `last_review` column.

```python
# compute and store the answer in the variable 'last_review_before_2015', then display it
count = 0
for idx in range(num_rows):
    if cell(idx, "last_review") == None:
        count += 0
    elif int(cell(idx, "last_review")[0:4]) <= 2015:
        count += 1
last_review_before_2015 = count
last_review_before_2015
```

Output:
```
1672
```

```python
grader.check("q9")
```

Output:
```
q9 results: All test cases passed!
```

### Function 2: `avg_price_per_room_type(room_type, neighborhood)`

We require you to complete the below function to answer the next several questions (this is a **requirement**, and you will **lose points** if you do not implement this function).

```python
def avg_price_per_room_type(room_type, neighborhood):
    total_price = 0
    count = 0
    prices_list = []
    for idx in range(num_rows):
        if cell(idx, "room_type") == room_type and cell(idx, "neighborhood") == neighborhood and cell(idx, "price") != None:
            total_price += cell(idx, "price")
            prices_list.append(cell(idx, "price"))
    if len(prices_list) == 0:
        return None
    return sum(prices_list)/len(prices_list)
```

**Question 10:** What is the  **average** cost of a *Private room* in the`neighborhood` *Little Neck*?

You **must** call the `avg_price_per_room_type` function to answer this question.

**Hint:** To help you debug your code in case you run into any bugs, we have reproduced in the cell below, **all** the rows in the dataset from the `neighborhood` *Little Neck*. If you run into bugs with `avg_price_per_room_type`, it is recommended that you go through your code and verify that it does what it is supposed to, for this tiny dataset.

room_id|name|host_id|host_name|neighborhood_group|neighborhood|latitude|longitude|room_type|price|minimum_nights|number_of_reviews|last_review|reviews_per_month|calculated_host_listings_count|availability_365
------|------|------|------|------|------|------|------|------|------|------|------|------|------|------|------|
20227428|Douglaston Apartment Room A|18996093|Leonard|Queens|Little Neck|40.75794000000001|-73.72955999999998|Private room|45|1|12|2019-06-22|0.55|5|133
21025083|Douglaston (apt 2) Room one\n(Largest room)|18996093|Leonard|Queens|Little Neck|40.75777|-73.72949|Private room|50|1|6|2018-12-16|0.31|5|94
30325639|Cozy shared studio in a safe neighborhood|21495656|Ramy|Queens|Little Neck|40.76212|-73.71928|Shared room|32|3|1|2018-12-04|0.14|1|88
31553066|Near major transportation|41090359|Abi|Queens|Little Neck|40.77122|-73.738|Private room|100|1|0|||1|88
35515780|30-min to Manhattan Quiet Big House in Great Neck|31859704|Vincent|Queens|Little Neck|40.77444000000001|-73.73373000000002|Entire home/apt|149|3|0|||1|3

```python
# compute and store the answer in the variable 'pvt_room_little_neck', then display it
pvt_room_little_neck = avg_price_per_room_type("Private room", "Little Neck")
pvt_room_little_neck
```

Output:
```
65.0
```

```python
grader.check("q10")
```

Output:
```
q10 results: All test cases passed!
```

**Question 11:** How much more expensive is a *Entire home/apt* than a *Private room* in the `neighborhood` *Astoria*?

You **must** call the `avg_price_per_room_type` function to answer this question.

```python
# compute and store the answer in the variable 'home_pvt_room_astoria_diff', then display it
home_pvt_room_astoria_diff = avg_price_per_room_type("Entire home/apt", "Astoria") - avg_price_per_room_type("Private room", "Astoria")
home_pvt_room_astoria_diff
```

Output:
```
46.02133741379495
```

```python
grader.check("q11")
```

Output:
```
q11 results: All test cases passed!
```

### Function 3: `find_prices_within(lat_min, lat_max, long_min, long_max)` 

We require you to complete the below function to answer the next several questions (this is a **requirement**, and you will **lose points** if you do not implement this function).

```python
def find_prices_within(lat_min, lat_max, long_min, long_max):
    price_list = []
    for idx in range(num_rows):
        lat = cell(idx, "latitude")
        long = cell(idx, "longitude")
        if lat_min <= lat <= lat_max and long_min <= long <= long_max:
            price_list.append(cell(idx, "price"))
    return price_list
```

**Question 12:** What is the **lowest** `price` room near *NYU* (`40.729 <= latitude <= 40.73, -74.01 <= longitude <= -74.00`)?

You **must** call the `find_prices_within` function to answer this question.

```python
# compute and store the answer in the variable 'min_price_nyu', then display it
min_price_nyu = min(find_prices_within(40.729, 40.73, -74.01, -74.00))
min_price_nyu
```

Output:
```
75
```

```python
grader.check("q12")
```

Output:
```
q12 results: All test cases passed!
```

### Function 4: `median(items)` 

We require you to complete the below function to answer the next several questions (this is a **requirement**, and you will **lose points** if you do not implement this function).

```python
def median(items):
    sorted_list = sorted(items)
    list_len = len(sorted_list)
    if list_len%2 == 1:
        return sorted_list[int((list_len/2) - 0.5)]
    else:
        first_middle = sorted_list[int((list_len/2) - 1)]
        second_middle = sorted_list[int(list_len/2)]
        return (first_middle + second_middle) / 2
```

**Question 13:** What is the **median** `price` of the rooms near *Columbia University* (`40.79 <= latitude <= 40.80, -73.96 <= longitude <= -73.95`)?

You **must** call the `find_prices_within` function to answer this question.

```python
# compute and store the answer in the variable 'median_price_columbia', then display it
median_price_columbia = median(find_prices_within(40.79, 40.80, -73.96, -73.95))
median_price_columbia
```

Output:
```
100
```

```python
grader.check("q13")
```

Output:
```
q13 results: All test cases passed!
```

**Question 14:** What **percentage** of rooms near *Rockefeller Center* (`40.749 <= latitude <= 40.75, -73.98 <= longitude <= -73.97`) have a `price` **more than** *100*?

Your answer **must** be a *float* value between *0* and *100*. You **must** call the `find_prices_within` function to answer this question.

```python
# compute and store the answer in the variable 'pct_price_over_hundred', then display it
over_hundred = []
for idx in range(num_rows):
    lat = cell(idx, "latitude")
    long = cell(idx, "longitude")
    if 40.749 <= lat <= 40.75 and -73.98 <= long <= -73.97 and cell(idx, "price") > 100:
        over_hundred.append(cell(idx, "price"))
top_num = len(over_hundred)
bottom_num = len(find_prices_within(40.749, 40.75, -73.98, -73.97))
pct_price_over_hundred = (top_num/bottom_num)*100
pct_price_over_hundred
```

Output:
```
93.10344827586206
```

```python
grader.check("q14")
```

Output:
```
q14 results: All test cases passed!
```

### Function 5: `review_avail_ratio(neighborhood)` 

We require you to complete the below function to answer the next several questions (this is a **requirement**, and you will **lose points** if you do not implement this function).

You should ignore rooms that have `availability_365` data of 0. You should also ignore rooms for which the ratio cannot be computed due to missing data.

1.  The denominator is the availability of a room (`availability_365` column). The numerator is the number of reviews of a room (`number_of_reviews column`). 
2.  Be careful! You need to compute the ratio for each room in the given neighborhood, then take the average of those ratios. Simply dividing the sum of reviews by the sum of availability will calculate the wrong answer.


```python
def review_avail_ratio(neighborhood):
    """
    review_avail_ratio(neighborhood) returns the average of the ratios of 
    number of reviews to availability of all rooms in the `neighborhood`
    """
    pass # replace with your code
    # TODO: you should **ignore** rooms that have `availability_365` data of 0
    # TODO: you should **ignore** rooms for which the ratio cannot be computed due to missing data
    # Hint: the numerator is the number of reviews of a room (`number_of_reviews column`)
    # Hint: the denominator is the availability of a room (`availability_365` column)
    # Hint: note that you need to compute the average of the ratios, **not** the ratio of the averages.
    #       you must compute the ratio for each room in the `neighborhood`, then take the average of those ratios.
    #       simply dividing the sum of reviews by the sum of availability will calculate the wrong answer.
    all_ratios = []
    for idx in range(num_rows):
        revs = cell(idx, "number_of_reviews")
        avail = cell(idx, "availability_365")
        if cell(idx, "neighborhood") == neighborhood:
            if cell(idx, "availability_365") == 0:
                continue
            all_ratios.append(revs/avail)
    if len(all_ratios) == 0:
        return None
    final_avg = sum(all_ratios)/len(all_ratios)
    return final_avg
```

**Question 15:** What is the **average of the ratios** of the `number_of_reviews` to `availability_365` in the `neighborhood` *Bushwick*?

You **must** call the `review_avail_ratio` function to answer this question.

```python
# compute and store the answer in the variable 'bushwick_avg_ratio', then display it
bushwick_avg_ratio = review_avail_ratio("Bushwick")
bushwick_avg_ratio
```

Output:
```
0.9856466156705735
```

```python
grader.check("q15")
```

Output:
```
q15 results: All test cases passed!
```

**Question 16:** What is the **average of the ratios** of the `number_of_reviews` to `availability_365` in the `neighborhood` *Manhattan Beach*?

You **must** call the `review_avail_ratio` function to answer this question.

```python
# compute and store the answer in the variable 'manhattan_beach_avg_ratio', then display it
manhattan_beach_avg_ratio = review_avail_ratio("Manhattan Beach")
manhattan_beach_avg_ratio
```

Output:
```
0.27323293295076073
```

```python
grader.check("q16")
```

Output:
```
q16 results: All test cases passed!
```

**Question 17:** Which `neighborhood` in the `neighborhood_group` *Staten Island* has the **highest average of ratios** of the `number_of_reviews` to `availability_365`?

**Clarification:** Don't worry about it if this cell takes around 10 seconds to run, that is expected. If it takes much longer (i.e., more than 30 seconds), you **must** optimize your code. Attend office hours if you are unable to get your code to run faster.

**Hint:** You do not need to compute the average of ratios for the **same** `neighborhood` more than once. Make a list of the **unique** neighborhoods in *Staten Island* first, and then find the **highest average of ratios** among those `neighborhoods`.

```python
# compute and store the answer in the variable 'max_nbhd_staten_island', then display it
max_nbhd_staten_island = ""
staten_nbs = []
numbers_rat = []
great_num = 0
for idx in range(num_rows):
    neigh = cell(idx, "neighborhood")
    if cell(idx, "neighborhood_group") == "Staten Island":
        if neigh in staten_nbs:
            continue
        else:
            staten_nbs.append(neigh)
            numbers_rat.append(review_avail_ratio(neigh))

for idx in range(len(staten_nbs)):
    if numbers_rat[idx] == None:
        continue
    if numbers_rat[idx] > great_num:
        great_num = numbers_rat[idx]
        max_nbhd_staten_island = staten_nbs[idx]
max_nbhd_staten_island

```

Output:
```
'Arden Heights'
```

```python
grader.check("q17")
```

Output:
```
q17 results: All test cases passed!
```

### Function 6: `secondary_word_in_found_rooms(find_room_word, secondary_word)`


We require you to complete the below function to answer the next several questions (this is a **requirement**, and you will **lose points** if you do not implement this function).

`secondary_word_in_found_rooms` function definition **must** invoke the function `find_room_names`. **We'll manually deduct points** if you don't use `find_room_names`.

```python
def secondary_word_in_found_rooms(find_room_word, secondary_word):
    """
    secondary_word_in_found_rooms(find_room_word, secondary_word) returns 
    the percentage of names containing the word `find_room_word` (case insensitive match)
    that also contains the word `secondary_word` (case insensitive match).
    """    
    pass # replace with your code
    # Hint: the denominator is the number of rooms with `find_room_word` in their name.
    # Hint: the numerator is the number of rooms that have **both** `find_room_word` and `secondary_word` in their name.
    first_list = find_room_names(find_room_word)
    sec_list = find_room_names(secondary_word)
    count = 0
    for idx in range(num_rows):
        name = cell(idx, "name")
        if name in first_list:
            if name in sec_list:
                count += 1
    denom = len(first_list)
    return (count/denom)*100
```

**Question 18:** What **percentage** of rooms whose `names` contain the word `"downtown"` **also** contain the word `"spacious"`?

Your answer **must** be a *float* value between *0* and *100*. You **must** call the `secondary_word_in_found_rooms` function to answer this question.

```python
# compute and store the answer in the variable 'downtown_and_spacious', then display it
downtown_and_spacious = secondary_word_in_found_rooms("downtown", "spacious")
downtown_and_spacious
```

Output:
```
5.541561712846348
```

```python
grader.check("q18")
```

Output:
```
q18 results: All test cases passed!
```

**Question 19:** What **percentage** of rooms whose `names` contain the word `"university"` **also** contain the word `"studio"`?

Your answer **must** be a *float* value between *0* and *100*. You **must** call the `secondary_word_in_found_rooms` function to answer this question.

```python
# compute and store the answer in the variable 'university_and_studio', then display it
university_and_studio = secondary_word_in_found_rooms("university", "studio")
university_and_studio
```

Output:
```
6.837606837606838
```

```python
grader.check("q19")
```

Output:
```
q19 results: All test cases passed!
```

**Question 20:** On a trip to NYC, you need to stay for *3* days in *Queens*, and then *4* days in *Brooklyn*. What is the **minimum** amount of money you need to spend on this trip?

Note that:
1. The price of each room is for each day.
2. The total cost = (lowest price in Queens) * 3 + (lowest price in Brooklyn) * 4.
3. You'll need to **skip** those rooms that don't have enough availability.


**Hint:** You might want to define a helper function to compute the **minimum** daily `price` of a room in a given `neighborhood_group` among rooms whose `availability_365` is equal to or larger than the number of days one will be staying in that neighborhood group.

```python
# compute and store the answer in the variable 'min_cost_trip', then display it
min_price_q = None
min_price_b = None
for idx in range(num_rows):
    if cell(idx, "neighborhood_group") == "Queens":
        if cell(idx, "minimum_nights") <= 3:
            if  min_price_q == None or cell(idx, "price") < min_price_q:
                min_price_q = cell(idx, "price")
for idx in range(num_rows):
    if cell(idx, "neighborhood_group") == "Brooklyn":
        if cell(idx, "minimum_nights") <= 4:
            if  min_price_b == None or cell(idx, "price") < min_price_b:
                min_price_b = cell(idx, "price")
min_cost_trip = (min_price_q*3) + (min_price_b*4)
min_cost_trip
```

Output:
```
30
```

```python
grader.check("q20")
```

Output:
```
q20 results: All test cases passed!
```

## Submission

Make sure you have run all cells in your notebook in order before running the cell below, so that all images/graphs appear in the output. The cell below will generate a zip file for you to submit. **Please save before exporting!**

**SUBMISSION INSTRUCTIONS**: 1. **Save** the notebook file **now (before you run the next cell of code)**. 2. **Upload** the zipfile to Gradescope. 3. Check **Gradescope otter** results as soon as the auto-grader execution gets completed. Don't worry about the score showing up as -/100.0. You only need to check that the test cases passed.

```python
# Save your notebook first, then run this cell to export your submission.
grader.export(pdf=False, run_tests=True)
```

Output:
```
Running your submission against local test cases...


Your submission received the following results when run against available test cases:

    q1 results: All test cases passed!

    q2 results: All test cases passed!

    q3 results: All test cases passed!

    q4 results: All test cases passed!

    q5 results: All test cases passed!

    q6 results: All test cases passed!

    q7 results: All test cases passed!

    q8 results: All test cases passed!

    q9 results: All test cases passed!

    q10 results: All test cases passed!

    q11 results: All test cases passed!

    q12 results: All test cases passed!

    q13 results: All test cases passed!

    q14 results: All test cases passed!

    q15 results: All test cases passed!

    q16 results: All test cases passed!

    q17 results: All test cases passed!

    q18 results: All test cases passed!

    q19 results: All test cases passed!

    q20 results: All test cases passed!
<IPython.core.display.HTML object>
```

Tests for q1 (points: 5)
```
>>> p6_test.check("q1", neighborhood_groups)
True
```


Tests for q10 (points: 5)
```
>>> p6_test.check("q10", pvt_room_little_neck)
True
```


Tests for q11 (points: 5)
```
>>> p6_test.check("q11", home_pvt_room_astoria_diff)
True
```


Tests for q12 (points: 5)
```
>>> p6_test.check("q12", min_price_nyu)
True
```


Tests for q13 (points: 5)
```
>>> p6_test.check("q13", median_price_columbia)
True
```


Tests for q14 (points: 5)
```
>>> p6_test.check("q14", pct_price_over_hundred)
True
```


Tests for q15 (points: 5)
```
>>> p6_test.check("q15", bushwick_avg_ratio)
True
```


Tests for q16 (points: 5)
```
>>> p6_test.check("q16", manhattan_beach_avg_ratio)
True
```


Tests for q17 (points: 5)
```
>>> p6_test.check("q17", max_nbhd_staten_island)
True
```


Tests for q18 (points: 5)
```
>>> p6_test.check("q18", downtown_and_spacious)
True
```


Tests for q19 (points: 5)
```
>>> p6_test.check("q19", university_and_studio)
True
```


Tests for q2 (points: 5)
```
>>> p6_test.check("q2", avg_price)
True
```


Tests for q20 (points: 5)
```
>>> p6_test.check("q20", min_cost_trip)
True
```


Tests for q3 (points: 5)
```
>>> p6_test.check("q3", count_soho)
True
```


Tests for q4 (points: 5)
```
>>> p6_test.check("q4", rooms_contain_cbg)
True
```


Tests for q5 (points: 5)
```
>>> p6_test.check("q5", rooms_contain_cinema_film)
True
```


Tests for q6 (points: 5)
```
>>> p6_test.check("q6", funnest_room)
True
```


Tests for q7 (points: 5)
```
>>> p6_test.check("q7", no_cost_rooms)
True
```


Tests for q8 (points: 5)
```
>>> p6_test.check("q8", pricey_neighborhoods)
True
```


Tests for q9 (points: 5)
```
>>> p6_test.check("q9", last_review_before_2015)
True
```
